---
title       : Version control
description : In this chapter, you'll learn about version control, Git, and how to use Git within RStudio. You'll start with the basics, like how to turn on version control using Git in an existing directory. Then, you'll see how to track changes to files over time, explore your commit history, get rid of unwanted changes, and how to push changes to remote repositories on GitHub.
attachments :


--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:db01ff97de
## What is version control?

--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:bda045cc6e
## Reproducibility and version control


--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:a27237d666
## Introduction to Git

<strong>Follow the directions in the console to open the R Markdown file:</strong>

```Type 'git()' and hit Enter to get started!```

You can use the RStudio platform just like you would on your home computer.

The markdown file contains the instuctions to answer the multiple choice question below.

<strong>Q1. What is Git?</strong>

*** =instructions
- Information manager from hell 
- A particular implementation of version control
- All of the above 
*** =hint
Check out the instructions in the .Rmd file to answer the multiple choice question.

*** =sct
```{r,eval=FALSE}
msg1 <- "Try again."
msg2 <- "Not quite."
msg3 <- "Great work!"
test_mc(3, feedback_msgs = c(msg1, msg2, msg3))
```
*** =attachments
.Rprofile: https://s3.amazonaws.com/assets.datacamp.com/production/course_1156/datasets/lab1_ex1.Rprofile
.init.R: https://s3.amazonaws.com/assets.datacamp.com/production/course_1156/datasets/lab1_ex1.init.R
lab1_ex1.Rmd: https://raw.githubusercontent.com/pzumbo/git/master/datasets/abc_git_ex1.Rmd
lab.css: https://s3.amazonaws.com/assets.datacamp.com/production/course_1156/datasets/lab.css


--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:bbca442764
## Version control in RStudio IDE


--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:bf83f3a9b7
## Initializing Git

--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:b8a90154e0
## What is a commit?


--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:ad6cfbfcca
## Git icons in RStudio

--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:cdf20e6403
## Commit history



--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:d8ce20fc5c
## Using checkout

--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:e83a93f2a8
## Undo commited changes: checkout


--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:941c11c480
## Undo  changes


--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:adbd3f800d
## Pulling & pushing


--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:9c1c3faf15
## Introduction to GitHub

--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:f6812b14bb
## Cloning exists repos

--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:5705dfbac5
## Working with branches

--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:211e19754b
## Summary






