---
title       : Introduction
description : 
attachments :


--- type:RStudioMultipleChoiceExercise lang:r xp:50 skills:1 key:db01ff97de
## Introduction

-----------------------------------------------------------------

**Which of the following is *not* a benefit of the RStudio IDE?**

*** =instructions
- Integrated user interface
- Contains a text editor
- Code completion
- Corrects bad code  

*** =sct
```{r,eval=FALSE}
msg1 <- "Try again."
msg2 <- "Not quite."
msg3 <- "Wrong."
msg4 <- "Correct!"
test_mc(4, feedback_msgs = c(msg1, msg2, msg3, msg4))
```
